#### Homework-1
##### Total number of points: 50
#### Due date: January 28th, 2022

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on basics of Python. The Python notebooks shared will be helpful to solve these problems. 

Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.

#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

#### Q1[5 points]. Assign the value for `x` to `150`. Set the value for `y` to `15` times `x` , and set the value for `z` to `y` divided by `x^2` (square of `x`).

<!-- grade_id: Q1-sol, points: 0, grade: false, solution: true -->
```python

#Assign x; y; z

x = 150
y = 15 * x
z = y / x**2

```

<!-- grade_id: 1, points: 3, grade: true, solution: false -->
```python
#Test cell-1
#DO NOT MODIFY/DELETE THIS CELL 

assert x==150
assert y==2250
assert z==0.1


```

<!-- grade_id: 2, points: 2, grade: true, solution: false -->
```python
#Test cell-2 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q2-1 [5 points]: 

- Create an empty list `li1`
- Create a list `li2` with integers 1, 2, 3
- Create a nested list `li3` with `li1` and `li2`
- Create a list `li4` with only `4` repeated 3 times.
- Create a list `li5` with your firstName and lastName as two separate strings -- for example, li5 that we build will be ['lydia', 'manikonda']


<!-- grade_id: Q2-sol, points: 0, grade: false, solution: true -->
```python
#Need to create li1, li2, li3, li4, li5 according to the instructions

li1 = []
li2 = [1,2,3]
li3 = [li1,li2]
li4 = [4,4,4]
li5 = ['Brooks','Bradbury']
```

<!-- grade_id: 3, points: 4, grade: true, solution: false -->
```python
#Test cell-3
#DO NOT MODIFY/DELETE THIS CELL 
assert len(li1) == 0
assert set(li2) == {1,2,3}
assert len(li3) == 2
assert len(set(li4)) == 1
assert len(li5)==2
```

<!-- grade_id: 4, points: 1, grade: true, solution: false -->
```python
#Test cell-4 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q2-2 [5 points]: 

###### Note that all are integers in this list. Please follow the exact order of execution here below from first line to the last line
- Create a list `li1` with values in an order -- 10, -10, 20, -20 
- Append a new element `30` to list `li1`  
- Insert a new element `-30` at position `2`
- Insert a new element `15` at position `20` (Yes it is 20 -- check what happens)

<!-- grade_id: Q2-2-Sol, points: 0, grade: false, solution: true -->
```python
#Need to create li1 according to the instructions above

li1 = [10,-10,20,-20]
li1.append(30)
li1.insert(2,-30)
li1.insert(20,15)

```

<!-- grade_id: 5, points: 1, grade: true, solution: false -->
```python
#Test cell-5 
#DO NOT MODIFY/DELETE THIS CELL
assert len(li1)==7
```

<!-- grade_id: 6, points: 4, grade: true, solution: false -->
```python
#Test cell-6 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q3-1 [5 points] 

###### Create a dictionary `d1` with 5 keys 1, 2, 3, 4, 5 and their corresponding values as 2, 4, 6, 8, 10 respectively. 

<!-- grade_id: Q3-Sol, points: 0, grade: false, solution: true -->
```python
#Creating the dictionary d1

d1 = {1:2,2:4,3:6,4:8,5:10}
```

<!-- grade_id: 7, points: 3, grade: true, solution: false -->
```python
#Test cell-7
#DO NOT MODIFY/DELETE THIS CELL
assert len(d1)==5
assert d1[3]==6
assert d1[5]==10
```

<!-- grade_id: 8, points: 2, grade: true, solution: false -->
```python
#Test cell-8 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q3-2 [5 points] 

###### Create a dictionary `d2` with 5 keys 1, 2, 3, 4, 5 and their corresponding values as 1, 4, 9, 16, 25 respectively. 
###### Add a new key-value pair to `d2` with key as `6` and value as `36`

<!-- grade_id: Q3-2-Sol, points: 0, grade: false, solution: true -->
```python
#Creating the dictionary d2 and add a new key-value pair

d2 = {1:1,2:4,3:9,4:16,5:25}
d2[6] = 36
```

<!-- grade_id: 9, points: 2, grade: true, solution: false -->
```python
#Test cell-9
#DO NOT MODIFY/DELETE THIS CELL

assert len(d2)==6
```

<!-- grade_id: 10, points: 3, grade: true, solution: false -->
```python
#Test cell-10 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q3-3 [5 points].

- Create a list `l31` with numbers 1, 2, 3, 4 in the same order as listed here
- Create another list `l32` with strings 'a', 'b', 'c', 'd' in the same order as listed here
- Create a dictionary `d3` where length of `l31` is the key and `l32` is the value
- Create a list `l33` that contains keys of dictionary `d3`

<!-- grade_id: Q3-3-Sol, points: 0, grade: false, solution: true -->
```python
#Creating l31 and l32 and follow the instructions to create dictionary d3 

l31 = [1,2,3,4]
l32 = ['a','b','c','d']
d3 = {len(l31):l32}
l33 = [d3.keys()]

```

<!-- grade_id: 11, points: 4, grade: true, solution: false -->
```python
#Test cell-11
#DO NOT MODIFY/DELETE THIS CELL
assert len(l31)==4
assert set(l32)=={'a', 'b','c', 'd'}
assert len(d3)==1
assert len(l33)==1
```

<!-- grade_id: 12, points: 1, grade: true, solution: false -->
```python
#Test cell-12 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q4-1 [5 points]. 

- Create a set `S1` with values 1, 2, 3, 4 in the same order as stated. 
- Assign the variable `Slen` as the length of this set `S1` 
- Create another set `S2` with values 5, 6, 7, 8 in the same order as stated. 
- Create a new set `S3` by merging sets `S1` and `S2`

<!-- grade_id: Q4-Sol, points: 0, grade: false, solution: true -->
```python
#Please follow the above instructions to create S1, S2, S3


S1 = {1,2,3,4}
Slen = len(S1)
S2 = {5,6,7,8}
S3 = set.union(S1,S2)

```

<!-- grade_id: 13, points: 3, grade: true, solution: false -->
```python
#Test cell-13
#DO NOT MODIFY/DELETE THIS CELL
assert Slen==4
assert 1 in S1
assert S3=={1, 2, 3, 4, 5, 6, 7, 8}
```

<!-- grade_id: 14, points: 2, grade: true, solution: false -->
```python
#Test cell-14 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q4-2 [5 points]

- Create a set `S4` with values 1, 2, 3, 4, 5, 6, 7, 8 passed as strings 
- Add another string `'10'` to `S4`
- Convert each element in S4 which is a string to integer and save this new set as `S5`

<!-- grade_id: Q4-2-Sol, points: 0, grade: false, solution: true -->
```python
S4 = set(['1','2','3','4','5','6','7','8'])
S4.add('10')
S5 = set()
for i in range(len(S4)-1):
  S5.add(int(i+1))
S5.add(10)
```

<!-- grade_id: 15, points: 5, grade: true, solution: false -->
```python
#Test cell-15
#DO NOT MODIFY/DELETE THIS CELL
assert len(set(S4))==9
assert S5=={1, 2, 3, 4, 5, 6, 7, 8, 10}
```

#### Q5 [5 points]. Write a function named `PrintNums()` that takes a number `num` (where `num` is >0) and returns a list of numbers from 0 to `num` (exclusive). If number `num`<=0 return `[-1]`. 

- Please make sure you consider the edgecases so that you will receive full points 

<!-- grade_id: Q5-Sol, points: 0, grade: false, solution: true -->
```python
def PrintNums(num):
  numlist = []
  if num > 0:
    for i in range(num):
      numlist.append(i)
    return numlist
  if num <= 0:
    return [-1]

```

<!-- grade_id: 16, points: 3, grade: true, solution: false -->
```python
#Test cell-16
#DO NOT MODIFY/DELETE THIS CELL
assert PrintNums(3)==[0, 1, 2]
assert PrintNums(4)==[0, 1, 2, 3]
assert PrintNums(6)==[0, 1, 2, 3, 4, 5]
```

<!-- grade_id: 17, points: 2, grade: true, solution: false -->
```python
#Test cell-17 HIDDEN TESTS
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q6 [5 points]. Write a function named 'GreaterOrLesser()' that takes two numbers `num1` and `num2`. Note that both `num1` and `num2` are `> 0`.  This function should: 
- return [1] if `num1` > `num2`
- return [0] if `num1` = `num2`
- return [-1] if `num1` < `num2`


<!-- grade_id: Q6-Sol, points: 0, grade: false, solution: true -->
```python
def GreaterOrLesser(num1,num2):
  if num1 > num2:
    return [1]
  elif num1 < num2:
    return [-1]
  else:
    return [0]

```

<!-- grade_id: 18, points: 5, grade: true, solution: false -->
```python
#Test cell-18
#DO NOT MODIFY/DELETE THIS CELL
assert GreaterOrLesser(0,1)==[-1]
assert GreaterOrLesser(num2=-1,num1=1)==[1]
assert GreaterOrLesser(num1=0,num2=0)==[0]

```