#### [Spring 2022] Homework-2 
##### Total number of points: 70
#### Due date: February 4th, 2022

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on basics of Python. The Python notebooks shared will be helpful to solve these problems. 

Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.

#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

#### Q1 [10 points]. Write a function that takes a number and checks if it is a prime number. Return `True` if its a prime number or else `False`. 


<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python
!pip install nose

def isPrimeNum(num):
  if num % 2 == 1 and num // 2 >= 1:
    return True
  else:
    return False
```

Output:
```
Collecting nose
  Downloading nose-1.3.7-py3-none-any.whl (154 kB)
[?25l[K     |██▏                             | 10 kB 20.0 MB/s eta 0:00:01[K     |████▎                           | 20 kB 20.8 MB/s eta 0:00:01[K     |██████▍                         | 30 kB 6.5 MB/s eta 0:00:01[K     |████████▌                       | 40 kB 8.2 MB/s eta 0:00:01[K     |██████████▋                     | 51 kB 7.3 MB/s eta 0:00:01[K     |████████████▊                   | 61 kB 7.2 MB/s eta 0:00:01[K     |██████████████▉                 | 71 kB 7.2 MB/s eta 0:00:01[K     |█████████████████               | 81 kB 8.0 MB/s eta 0:00:01[K     |███████████████████             | 92 kB 6.3 MB/s eta 0:00:01[K     |█████████████████████▏          | 102 kB 6.8 MB/s eta 0:00:01[K     |███████████████████████▎        | 112 kB 6.8 MB/s eta 0:00:01[K     |█████████████████████████▍      | 122 kB 6.8 MB/s eta 0:00:01[K     |███████████████████████████▌    | 133 kB 6.8 MB/s eta 0:00:01[K     |█
```

<!-- grade_id: 1, points: 6, grade: true, solution: false -->
```python
#Test cell-1 [6 pts]
#DO NOT MODIFY/DELETE THIS CELL 

from nose.tools import assert_equal
'''case-1'''
number=1
assert isPrimeNum(number)==False

'''case-2'''
number = 10
assert isPrimeNum(number)==False

'''case-3'''
number = 97
assert_equal(isPrimeNum(number), True)

```

<!-- grade_id: 2, points: 4, grade: true, solution: false -->
```python
#Test cell-2 [4 pts]
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q2 [5 points]. import the numpy package as np. 

<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
import numpy as np
```

<!-- grade_id: 3, points: 5, grade: true, solution: false -->
```python
#Test cell-3 [5 pts]
#DO NOT MODIFY/DELETE THIS CELL 
assert np.prod([1,2,3])==6
```

#### Q3 [5 points]. Create a numpy array called `arr3` of zeros of size 10 where the fifth value (or position) is 1. 


<!-- grade_id: q3-q, points: 0, grade: false, solution: true -->
```python
arr3 = np.zeros((1,10))
arr3 = np.insert(arr3,4,[1])

```

<!-- grade_id: 4, points: 3, grade: true, solution: false -->
```python
#Test cell-4 [3 pts]
#DO NOT MODIFY/DELETE THIS CELL 

assert isinstance(arr3, np.ndarray)
assert arr3[:4].all() == 0
assert arr3[4] == 1

```

<!-- grade_id: 5, points: 2, grade: true, solution: false -->
```python
#Test cell-5 [2 pts]
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q4 [5 points]. Create a numpy array `arr4` with elements 1,3,5,7,9 (in the same order)
#### Now create a new array `arr41` whose elements are twice the value of each element in `arr4` at the same index respectively. 


<!-- grade_id: q4-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
arr4 = np.array([1,3,5,7,9])
arr41 = np.multiply(arr4,2)

```

<!-- grade_id: 6, points: 4, grade: true, solution: false -->
```python
#Test cell-6 [4 pts]
#DO NOT MODIFY/DELETE THIS CELL 

assert isinstance(arr4, np.ndarray)
assert arr4[2] == 5
assert arr41[2] == 10
assert arr41[1] == 6

```

<!-- grade_id: 7, points: 1, grade: true, solution: false -->
```python
#Test cell-7 [1 pts]
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q5 [10 points]. Return the absolute difference between the element that occurs maximum number of times and element that occurs minimum number of times in a numpy array. 
- abs() is a function used to return absolute values of a given number. For example, abs(-3) = 3 and abs(3) = 3 
- if there are two (or more) elements occuring with same frequency and are occurring maximum number of times, choose the element that is smaller in value
- if there are two (or more) elements occuring with same frequency and are occurring minimum number of times, choose the element that is smaller in value


<!-- grade_id: q5-sol, points: 0, grade: false, solution: true -->
```python
from numpy.ma.core import shape

from scipy import stats

def diffMinMax(arr5):
  highfreq = stats.mode(arr5)
  mode = int(highfreq[0])
  leastfreq = np.unique(arr5, return_counts = True)
  leastfreq = np.asarray(leastfreq)
  leastfreqsort = leastfreq [ :, leastfreq[1].argsort()]
  leastval = leastfreqsort[0,0]
  diff = abs(highfreq - leastval)
  diff = int(diff[0])
  return diff

```

<!-- grade_id: 8, points: 10, grade: true, solution: false -->
```python
#Test cell-8 [10 pts]
#DO NOT MODIFY/DELETE THIS CELL 

arr3=np.array([1,0,1,0,2,1])
assert diffMinMax(arr3)==1

arr3=np.array([0])
assert diffMinMax(arr3)==0

arr3=np.array([1,2,3,1,2,3,1,1,0,0,4])
assert diffMinMax(arr3)==3

arr3=np.array([1,2,3,0])
assert diffMinMax(arr3)==0
```

```python
import pandas as pd

'''Make sure you run this cell before executing the cells below'''
```

Output:
```
'Make sure you run this cell before executing the cells below'
```

#### Q6 [5 points]. Please follow the instructions very carefully 
- Create a list 'l6' and assign 1, 2, 3, 4, 5 to this list
- Use l6 to create a series S6 with a default index


<!-- grade_id: q6-sol, points: 0, grade: false, solution: true -->
```python
l6 = [1,2,3,4,5]
S6 = pd.Series(l6)

```

<!-- grade_id: 9, points: 3, grade: true, solution: false -->
```python
#Test cell-9 [3 pts]
#DO NOT MODIFY/DELETE THIS CELL 

assert len(l6)==5
assert max(l6)==5
assert S6[3]==4

```

<!-- grade_id: 10, points: 2, grade: true, solution: false -->
```python
#Test cell-10 [2 pts]
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q7 [10 points]. This function takes a list `l7` as an input and creates 2 more lists (`l71` and `l72`) by manipulating `l7` as follows
-  l71 -- copy l7 into l71 and add the 1st element of l7 to the end of l71
- l72 -- copy l71 into l72 and add the mean value of l71 to the end of l72
- Now use `l7`, `l71` and `l72` to create a data frame `df` with corresponding column names as `first`, `second` and `third`
- return this dataframe `df`

<!-- grade_id: q7-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd

def createDF(l7):
  first = l7[0]
  l71 = list(l7)
  l71.insert(999,first)
  l72 = list(l71)
  mean_l71 = sum(l71) / len(l71)
  l72.insert(999,mean_l71)

  d = {'first':pd.Series(l7),'second':pd.Series(l71),'third':pd.Series(l72)}
  df = pd.DataFrame(d)
  return df

```

<!-- grade_id: 11, points: 5, grade: true, solution: false -->
```python
#Test cell-11 [5 pts]
#DO NOT MODIFY/DELETE THIS CELL 

'''case-1'''
list7=[2,4,6,8]
assert list(createDF(list7)['second'])[:5]==[2.0,4.0,6.0,8.0,2.0]

'''case-2'''
list7=[2,4,6,8]
assert list(createDF(list7)['first'])[:4]==[2,4,6,8]

```

<!-- grade_id: 12, points: 5, grade: true, solution: false -->
```python
#Test cell-12 [5 pts]
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q8 [10 points]. Read the .csv file from this url and load it into dataframe 'df'
- Using pandas dataframe operations, output the total number of rows with `survived`==1 and `who` is `woman` as `var81`.
- Using pandas dataframe operations, output the total number of rows with `survived`==1 and `who` is `child` as  `var82`. 


<!-- grade_id: q8-sol, points: 0, grade: false, solution: true -->
```python

url='https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/titanic.csv'
df = pd.read_csv(url)
# df.head()
var81 = df.loc[(df['survived'] == 1) & (df['who'] == 'woman')]
var82 = df.loc[(df['survived'] == 1) & (df['who'] == 'man')]
var82 = len(var82)
```

<!-- grade_id: 13, points: 6, grade: true, solution: false -->
```python
#Test cell-13 [6 pts]
#DO NOT MODIFY/DELETE THIS CELL 

assert df['age'].sum()==21205.17
assert df.isnull().sum().sum()==869
assert var82==88

```

<!-- grade_id: 14, points: 4, grade: true, solution: false -->
```python
#Test cell-14 [4 pts]
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q9 [10 points]. Read the .csv file from this url and load it into dataframe `df`
#### Using groupby operation over `embark_town` and `gender`, 
- assign `var9` as the total number of people survived (`survived` is 1), where they embarked (`embark_town`) is at `Southampton` and are female (`gender` is `female`). 


<!-- grade_id: q9-sol, points: 0, grade: false, solution: true -->
```python

url='https://raw.githubusercontent.com/Spring2021IntrotoML/Spring2021IntrotoML.github.io/main/Lectures/titanic.csv'
df = pd.read_csv(url)
df.head()
var9 = df.loc[(df['survived'] == 1) & (df['embark_town'] == 'Southampton') & (df['gender'] == 'female')]
var9 = len(var9)
```

<!-- grade_id: 15, points: 10, grade: true, solution: false -->
```python
#Test cell-15 [10 pts]
#DO NOT MODIFY/DELETE THIS CELL 

assert df['age'].sum()==21205.17
assert df.isnull().sum().sum()==869
assert len(df.columns)==15
assert var9==140
```