#### Homework-5
##### Total number of points: 70
#### Due date: March 4th, 2022

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on basics of Python. The Python notebooks shared will be helpful to solve these problems. 

Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.

#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

#### Q1 [10 points]. Please make sure this is correct as the following questions are dependent on this. 
- Load the boston data from sklearn library (library is already provided here below). 
- The given target labels (boston.target) are continuous -- convert them into discrete values -- 1 and 2 using this approach.
- all the values (v1) in boston.target should become 1 if v1>=5 and v1<23; 
- all the values (v2) in boston.target should become 2 if v2>=23 and v2<51;
- Note that these new/transformed labels 1 and 2 are **integers**
- Please save these new discrete values in the variable `y_true`

<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston

#Load the dataset and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
# YOUR CODE HERE

test1 = np.where((boston.target >= 5) & (boston.target < 23), 1, boston.target)
test2 = np.where((test1 >= 23) & (test1 < 51), 2, test1)

y_true = test2
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 1, points: 10, grade: true, solution: false -->
```python
#[10 points] Test cell-1 
#DO NOT MODIFY/DELETE THIS CELL 
assert set(y_true)=={1,2}
assert (list(y_true)).count(1)==312
assert (list(y_true)).count(0)==0
assert (list(y_true)).count(2)==194
```

#### Q2 [5 points]. Now split the data into training (80%) and testing data (20%) by using these variable names 
- `X_train`: Training feature columns
- `X_test`: Testing feature columns
- `y_train`: Training labels
- `y_test`: Testing labels
- with only parameters df, y_true and 'test_size'. `df` and `y_true` are initialized in the previous question. 

<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from sklearn.model_selection import train_test_split

X_train = pd.DataFrame()
y_train = pd.DataFrame()
X_test = pd.DataFrame()
y_test = pd.DataFrame()

X_train, X_test, y_train, y_test = train_test_split(df, y_true, test_size=0.2)

```

<!-- grade_id: 2, points: 10, grade: true, solution: false -->
```python
#[10 points] Test cell-2
#DO NOT MODIFY/DELETE THIS CELL 
assert len(X_train)==404
assert len(y_test)==102
assert len(y_train)==404
assert len(X_test)==102
```

#### Q3 [5 points]. Use the `df` in Q1 to perform these operations: 
<ol>
<li> Standardize the data using StandardScaler() function from sklearn library</li>
<li> Compute principal components using the fit_transform operation on the original dataset </li>
<li> Now find the number of principal components `n_components` that will retain not more than 75% of the information present in the original dataset </li>
</ol>

- Make sure you declare appropriate Python packages which are not provided by default
- Hint -- Use explained_variance_ratio_.cumsum() function we discussed in the class


<!-- grade_id: q3-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn.preprocessing import StandardScaler
from sklearn.decomposition import PCA

#Load the dataset and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)

scaler = StandardScaler()
scaled = scaler.fit_transform(df)
df_scaled = pd.DataFrame(scaled, columns=list(df.columns))
df_scaled.head()

pca = PCA()
pca.fit(df_scaled)
components = pca.explained_variance_ratio_.cumsum()

n_components=0

for i in range(len(components)):
  if components[i] < 0.75:
    n_components = n_components + 1

print('The number of components under threshhold is:',n_components)
```

Output:
```
The number of components under threshhold is: 4
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datas
```

<!-- grade_id: 3, points: 2.5, grade: true, solution: false -->
```python
#[2.5 points] Test cell-3
#DO NOT MODIFY/DELETE THIS CELL 
assert n_components!=0
```

<!-- grade_id: 4, points: 2.5, grade: true, solution: false -->
```python
#[2.5 points] Test cell-4 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q4 [10 points]. Given 3 different series `S1`, `S2`, `S3`
<ol>
<li>First, using these 3 series build a dataframe 'df' where the column name associated with each series is same as the series name such as `S1`, `S2` or `S3`.</li>
<li>Please note that 'df' doesn't contain any column representing class labels. </li>
<li>If you perform PCA operation, how many maximum principal components will you obtain?</li>
<li>No need to do pca but assign the value to variable `max_pcs`</li>
</ol>

<!-- grade_id: q4-sol, points: 0, grade: false, solution: true -->
```python

import pandas as pd
S1 = pd.Series([1,2,3,4]) #corresponding column name should be 'S1'
S2 = pd.Series([10, 20, 30, 40]) #corresponding column name should be 'S2'
S3 = pd.Series([2,4,6,8]) #corresponding column name should be 'S3'

df = pd.DataFrame({'S1': S1, 'S2': S2, 'S3': S3})

scaler = StandardScaler()
scaled = scaler.fit_transform(df)
df_scaled = pd.DataFrame(scaled, columns=list(df.columns))
df_scaled.head()

pca = PCA()
pca.fit(df_scaled)

pca.explained_variance_ratio_.cumsum()

#max_pcs=0 #Update this variable after building the dataframe #No need to perform PCA 

# YOUR CODE HERE



```

Output:
```
array([1., 1., 1.])
```

<!-- grade_id: 5, points: 7.5, grade: true, solution: false -->
```python
#[7.5 points] Test cell-5
#DO NOT MODIFY/DELETE THIS CELL 
assert set(df.columns)=={'S1', 'S2', 'S3'}
assert df['S3'].mean()==5.0
assert df['S1'].mean()==2.5
```

<!-- grade_id: 6, points: 2.5, grade: true, solution: false -->
```python
#[2.5 points] Test cell-6 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q5 [10 points]. Linear regression: Use the boston dataset loaded into the dataframe `df`
##### 1. **Without using train_test_split() function** use the 1st 400 rows (i.e., index 0,1,....399) where, 
- the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
- Use the remaining rows as testing data -- `test_X` and `test_y` for the features and labels respectively.
- Fit a linear regression line using the training data; then use it to predict labels for testing data as shown in the lecture notebook. Please use the default parameters when calling the linear regression function. 

##### 2. Measure the mean-squared error (MSE) 'mse_split1' using the predicted labels with testing labels.
- Round the `mse_split1` to 2 values after the decimal point. 
- Hint -- Check out the libraries given to guess which functions to use to compute MSE.


<!-- grade_id: q5-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn import datasets
from sklearn import linear_model
from sklearn.metrics import mean_squared_error
from sklearn.linear_model import LinearRegression
from sklearn.metrics import accuracy_score

#Load the dataset without class labels and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
df

# class label -- boston.target

train_X = pd.DataFrame(df[:400])
train_y = pd.DataFrame(boston.target[:400])
test_X = pd.DataFrame(df[400:])
test_y = pd.DataFrame(boston.target[400:])

lm = LinearRegression().fit(train_X, train_y)
y_pred = lm.predict(test_X)
mse_split1 = round(mean_squared_error(test_y, y_pred),2)
print(mse_split1)
print(len(train_y))
print(len(test_X))


```

Output:
```
37.89
400
106
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing
```

<!-- grade_id: 7, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-7
#DO NOT MODIFY/DELETE THIS CELL 
assert len(train_X)==400
assert len(test_y)==106
assert mse_split1==37.89

```

<!-- grade_id: 8, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-8 Hidden tests 
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q6 [10 points]. Repeat the above exercise (Q5) but with different training and testing splits using the `boston` dataset.

###### 1. **Without using train_test_split() function** use the last 253 rows (i.e., index 253,254,....505) where, 
- the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
- Use the remaining rows as testing data -- `test_X` and `test_y` for the features and labels respectively.
- Fit a linear regression line using the training data; then use it for predicting testing data as shown in the lecture notebook. Please use the default parameters when calling the linear regression function. 

###### 2. Measure the mean-squared error (MSE) 'mse_split2' using the predicted labels with testing labels.
- Round the `mse_split2` to 2 values after the decimal point. 
- Hint -- Check out the libraries given to guess which functions to use to compute MSE.


<!-- grade_id: q6-sol, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn import datasets
from sklearn import linear_model
from sklearn.metrics import mean_squared_error

#Load the dataset without class labels and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
#class label -- boston.target

train_X = pd.DataFrame(df[253:])
train_y = pd.DataFrame(boston.target[253:])
test_X = pd.DataFrame(df[:253])
test_y = pd.DataFrame(boston.target[:253])

lm = LinearRegression().fit(train_X, train_y)
y_pred = lm.predict(test_X)
mse_split2 = round(mean_squared_error(test_y, y_pred),2)
print(mse_split2)


```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 9, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-9  
#DO NOT MODIFY/DELETE THIS CELL 
assert len(train_X)==253
assert len(test_y)==253
assert mse_split2==27.22

```

<!-- grade_id: 10, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-10 Hidden tests 
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q7 [15 points]. We have loaded the boston dataset -- use the dataframe `df` to do: 
###### 1. Convert the target labels -- boston.target from continuous to discrete values -- 1, 2 and 3 using this approach.
- all the values (lets say, each value is represented by v1) in boston.target should become 1 if v1>=5 and v1<20; 
- all the values (lets say, each value is represented by v2) in boston.target should become 2 if v2>=20 and v2<35;
- all the values (lets say, each value is represented by v3) in boston.target should become 3 if v3>=35 and v3<51;

###### 2. **Without using train_test_split()** use the 1st 400 rows (i.e., index 0,1,....399) where, 
- the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
- Use the rest of the rows for testing data -- `test_X` and `test_y` for the features and groundtruth labels respectively.
- Fit a logistic regression with solver=`newton-cg`, C=1e5, multi_class=`multinomial`; then use it for predicting testing data as shown in the lecture notebook. If you encounter this warning, please ignore it -- "ConvergenceWarning: newton-cg failed to converge." 

###### 3. Measure the accuracy `acc_split3` using the predicted labels with groundtruth test labels `test_y`.
- Round the `acc_split3` to 2 values after the decimal point. 
- Hint -- Check out the libraries given to guess which functions to use to compute accuracy.

<!-- grade_id: q7-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn import datasets
from sklearn import linear_model
from sklearn.metrics import accuracy_score
from sklearn.linear_model import LogisticRegression

#Load the dataset without class labels and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
#class label -- boston.target

phase1 = np.where((boston.target >= 5) & (boston.target < 20), 1, boston.target)
phase2 = np.where((phase1 >= 20) & (phase1 < 35), 2, phase1)
phase3 = np.where((phase2 >= 35) & (phase2 < 51), 3, phase2)
targetlabels = phase3

train_X = pd.DataFrame(df[:400])
train_y = pd.DataFrame(targetlabels[:400])
test_X = pd.DataFrame(df[400:])
test_y = pd.DataFrame(targetlabels[400:])

logR = LogisticRegression(solver='newton-cg',C=1e5,multi_class='multinomial').fit(train_X, train_y)
y_pred1 = logR.predict(df)
acc_split3 = round(accuracy_score(test_y, y_pred1[400:]),2)
  
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 11, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-11 
#DO NOT MODIFY/DELETE THIS CELL 
assert len(train_X)==400
assert set(targetlabels)=={1,2,3}
```

<!-- grade_id: 12, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-12  
#DO NOT MODIFY/DELETE THIS CELL 
assert acc_split3==0.74
assert len(test_y)==106
```

<!-- grade_id: 13, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-13 Hidden tests 
#DO NOT MODIFY/DELETE THIS CELL 
```