#### Homework-7
##### Total number of points: 30
#### Due date: April 15, 2022

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.
 

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

```python
#Please run this cell to load the dataset. 
url = 'https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/dataset_Facebook.csv'

import pandas as pd
import math
df  = pd.read_csv(url)
df = df.dropna()
df.columns
```

Output:
```
Index(['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday',
       'Post Hour', 'Paid', 'Lifetime Post Total Reach',
       'Lifetime Post Total Impressions', 'Lifetime Engaged Users',
       'Lifetime Post Consumers', 'Lifetime Post Consumptions',
       'Lifetime Post Impressions by people who have liked your Page',
       'Lifetime Post reach by people who like your Page',
       'Lifetime People who have liked your Page and engaged with your post',
       'comment', 'like', 'share', 'Total Interactions'],
      dtype='object')
```

#### Q1[30 points]: We will use neural networks to model a regression problem. 

#### 1. Create a new dataframe `df_sub` that contains only these feature columns in `df` : 
- `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`,  `Total Interactions`

#### 2. Perform one-hot encoding on these features below in the dataframe `df_sub4` to create a new dataframe `df_OHE`
- `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`

#### 3. Perform normalization only on `Page total likes` and `Total Interactions` column in `df_OHE` using `minmax_scale` function 
- note that both these functions will output the same result

#### 4. Using `df_OHE` perform train_test_split operation to build training (`X_train`, `y_train`) and testing data (`X_test`, `y_test`).
- Use test_size=0.3, random_state=42 as the parameters for train_test_split() function. 
- Feature columns (X): Everything except `Total Interactions`
- CLass label column (y): `Total Interactions`

#### 5. Build a 3-layer neural network with `relu` as the 2 hidden layers' activation function (12 nodes, 8 nodes in each hidden layer) and `sigmoid` as the final layer's (1 node) activation function. 
- Use the same `adam` optimizer and `mean_squared_error` as the loss function. 
- `epochs`=150, `batch_size`=10
- since the class label is continuous, we use `metrics.MeanSquaredError()` as the metric (`from keras import metrics`) to measure the performance of our classifier captured as variable `mse` 

```python
#Please run this cell to install keras
!pip install keras
!pip install tensorflow
```

Output:
```
Requirement already satisfied: keras in /usr/local/lib/python3.7/dist-packages (2.8.0)
Requirement already satisfied: tensorflow in /usr/local/lib/python3.7/dist-packages (2.8.0)
Requirement already satisfied: typing-extensions>=3.6.6 in /usr/local/lib/python3.7/dist-packages (from tensorflow) (4.1.1)
Requirement already satisfied: wrapt>=1.11.0 in /usr/local/lib/python3.7/dist-packages (from tensorflow) (1.14.0)
Requirement already satisfied: google-pasta>=0.1.1 in /usr/local/lib/python3.7/dist-packages (from tensorflow) (0.2.0)
Requirement already satisfied: absl-py>=0.4.0 in /usr/local/lib/python3.7/dist-packages (from tensorflow) (1.0.0)
Requirement already satisfied: h5py>=2.9.0 in /usr/local/lib/python3.7/dist-packages (from tensorflow) (3.1.0)
Requirement already satisfied: tensorflow-io-gcs-filesystem>=0.23.1 in /usr/local/lib/python3.7/dist-packages (from tensorflow) (0.24.0)
Requirement already satisfied: gast>=0.2.1 in /usr/local/lib/python3.7/dist-packages (from tensorflow)
```

<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
# Libraries
from sklearn.preprocessing import MinMaxScaler
from sklearn.preprocessing import minmax_scale
from sklearn.model_selection import train_test_split
from keras.models import Sequential
from keras.layers import Dense
from keras import metrics

# Part 1: Select certain columns

df_sub = df[['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour', 'Paid', 'Total Interactions']] # Copy variable names in question description and paste
df_sub.head()

# Part 2: OHC

df_OHE = pd.get_dummies(data = df_sub, columns = ['Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour'])
df_OHE['Page total likes'] = df_sub['Page total likes'].values
df_OHE['Paid'] = df_sub['Paid'].values
df_OHE['Total Interactions'] = df_sub['Total Interactions'].values
df_OHE.head()

# Part 3: Normalizing
cols_to_norm = ['Page total likes', 'Total Interactions']
df_OHE[cols_to_norm] = minmax_scale(df_OHE[cols_to_norm])
df_OHE.head()

# Part 4: Train test

X = df_OHE.drop(columns = ['Total Interactions'])
y = df_OHE['Total Interactions']

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size = 0.3, random_state = 42)

# Part 5: Neural Network // Copied from Lecture 22

model = Sequential()
model.add(Dense(12, input_dim=50, activation='relu'))
model.add(Dense(8, activation='relu'))
model.add(Dense(1, activation='sigmoid'))

# compile the keras model with mean_squared_error and optimizer with a certain learning rate
#Adam optimization is a stochastic gradient descent method that is based on adaptive estimation of first-order and second-order moments
model.compile(loss='mse', optimizer='adam', metrics=['mean_squared_error'])

# fit the keras model on the dataset
model.fit(X_train, y_train, epochs=150, batch_size=10)

# evaluate the keras model
_, mse = model.evaluate(X_test, y_test)
print('MSE: %.4f' % (mse))
```

Output:
```
Epoch 1/150
35/35 [==============================] - 1s 3ms/step - loss: 0.2215 - mean_squared_error: 0.2215
Epoch 2/150
35/35 [==============================] - 0s 2ms/step - loss: 0.1808 - mean_squared_error: 0.1808
Epoch 3/150
35/35 [==============================] - 0s 3ms/step - loss: 0.1379 - mean_squared_error: 0.1379
Epoch 4/150
35/35 [==============================] - 0s 3ms/step - loss: 0.0906 - mean_squared_error: 0.0906
Epoch 5/150
35/35 [==============================] - 0s 2ms/step - loss: 0.0496 - mean_squared_error: 0.0496
Epoch 6/150
35/35 [==============================] - 0s 2ms/step - loss: 0.0234 - mean_squared_error: 0.0234
Epoch 7/150
35/35 [==============================] - 0s 3ms/step - loss: 0.0113 - mean_squared_error: 0.0113
Epoch 8/150
35/35 [==============================] - 0s 4ms/step - loss: 0.0073 - mean_squared_error: 0.0073
Epoch 9/150
35/35 [==============================] - 0s 2ms/step - loss: 0.0059 - mean_squared_error: 0.0059
Epoch 10/150
35/35 
```

<!-- grade_id: 1, points: 10, grade: true, solution: false -->
```python
##Cell-1 -- Do not modify this cell
assert (math.ceil(df_OHE['Page total likes'].sum())) == 357
assert (math.ceil(df_OHE['Total Interactions'].sum())) == 17
assert (math.ceil(df_OHE['Page total likes'].mean())) == 1
assert (math.ceil(df_OHE['Total Interactions'].mean())) == 1
```

<!-- grade_id: 2, points: 10, grade: true, solution: false -->
```python
##Cell-2 -- Do not modify this cell
assert len(X_train)==346
assert len(y_test)==149
assert mse<0.005
assert len(model.weights)==6
```

<!-- grade_id: 3, points: 10, grade: true, solution: false -->
```python
##Cell-3 -- Do not modify this cell
assert set(df_sub.columns)=={'Post Month', 'Category', 'Page total likes', 'Total Interactions', 'Post Hour', 'Type', 'Post Weekday', 'Paid'}
assert len(df_sub)==495
assert len(df_OHE.columns)==51
assert len(df_OHE)==495
```